Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Abraham Esquivel"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
Because by that we mean that the worst case scenario running time and the running time of any input is O(n^2)

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
For the $2^{n+1}=O(2^n)$ equation it is true because it doesn't matter if you add any value to the n because it will still be part of th O notation. However, with $2^{2n} = O(2^n)$, by developing the equation, the actual big O for $2^{2n}$ is $O(4^n)$ and 

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    
    """
    # YOUR CODE HERE
    
    gmax=0   #This variable refers to the global maximum value
    i = []   #This variable refers to the indices. So far it's an empy array
    for x in range (len (A)+1):#this line and the next one (line 19) represent the first and the last indices, respectively,
                                #that are going to be in the maximum subarray                              
        for y in range(len(A)+1):
            if A[x:y]: #this line sets the baoundary of the array which means that it won't try to get out from there
                cmax = sum(A[x:y]) #cmax refers to the maximum value so far which is the summ of the elements of the array
                if cmax>gmax: #here it makes a comparison between the current maximum with the global maximum
                    gmax=cmax #and if it turns out to be true, then the current maximum turns to be the new global maximum
                    i.append((x,y-1, gmax)) #this line adds to i the first index of the maximum subarray; 
                    #the last number of it and the value of the maximum subarray, which is represented by gmax
    print(i.pop())#because this program will keep adding elements to i,
                  #this line shows only the result of the last iteration
    

    
    raise NotImplementedError()
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
(1, 3, 2)
<ipython-input-37-5ddab6f2f574> in <module>()
----> 1 assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
      2 assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))

<ipython-input-36-09d536f813fe> in bruteforce_max_subarray(A)
     27 
     28 
---> 29     raise NotImplementedError()

NotImplementedError:
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    gmax=0   #This variable refers to the global maximum value
    i = []   #This variable refers to the indices. So far it's an empy array
    for x in range (len (A)+1):#this line and the next one (line 19) represent the first and the last indices, respectively,
                                #that are going to be in the maximum subarray                              
        for y in range(len(A)+1):
            if A[x:y]: #this line sets the baoundary of the array which means that it won't try to get out from there
                cmax = sum(A[x:y]) #cmax refers to the maximum value so far which is the summ of the elements of the array
                if cmax>gmax: #here it makes a comparison between the current maximum with the global maximum
                    gmax=cmax #and if it turns out to be true, then the current maximum turns to be the new global maximum
                    i.append((x,y-1, gmax)) #this line adds to i the first index of the maximum subarray; 
                    #the last number of it and the value of the maximum subarray, which is represented by gmax
    return(i.pop()) #because this program will keep adding elements to i,
                   #this line shows only the result of the last iteration
    raise NotImplementedError()
```

```python
bruteforce_max_subarray([13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7])
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort | $O(n^2)$ |
Selection sort | $O(n^2)$ |
Bubble sort |$O(n^2)$ | 
Finding maximum subarray |  |$\Theta(n)$

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```